// strings1.rs
// Make me compile without changing the function signature!
// Execute `rustlings hint strings1` for hints ;)

// I AM NOT DONE

fn main() {
    let answer: String = current_favorite_color();
    println!("My current favorite color is {}", answer.to_string());
}

fn current_favorite_color() -> String {
    "blue"
}
